<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

```python
%pylab inline
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

# Mean Squrare Error

In the lecture we considered the sample-mean estimator for the distribution mean.
Another estimator for the distribution mean is the min-max-mean estimator that takes the mean (average) of the smallest
and largest observed values. For example, for the sample {1, 2, 1, 5, 1},
the sample mean is (1+2+1+5+1)/5=2 while the min-max-mean is (1+5)/2=3.
In this problem we ask you to run a simulation that approximates
the mean squared error (MSE) of the two estimators for a uniform distribution.

Take a continuous uniform distribution between $a$ and $b$ - given as parameters.
Draw a 10-observation sample from this distribution, and calculate
the sample-mean and the min-max-mean. Repeat the experiment 1,000,000
times, and for each estimator calculate its average bias, and its
(Bessel-corrected) sample-variance over the 1,000,000 estimates. 
Take the square of the average bias plus the sample variance
as your MSE estimates.

Note that you can use the function numpy.random.uniform to generate random samples. You can calculate the mean and variance using the formulas presented in the lecture or via the function numpy.mean and numpy.var. To calculate the unbiased sample variance using function numpy.var, you will need to set ddof=1.

## Problem 1

For this problem, you are asked to write a function **Sample_Mean_MSE** that experimentally calculates the MSE of the sample-mean estimator for samples drawn from uniform distribution given the bounds $a$ and $b$.

<font  style="color:blue"> **Code:**</font>
```python
Sample_Mean_MSE(0,10)      
```


<font  style="color:magenta"> **Output**</font>
```
0.83433254206591267
```

```python
# modify this cell

def Sample_Mean_MSE(a,b):
    # inputs: bounds for uniform distribution a and b
    # sample size is 10
    # number of experiments is 1,000,000
    # output: MSE for sample mean estimator with sample size 10
    
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
# Check Function
assert abs(Sample_Mean_MSE(-5,5)-0.83333333333333) < 0.05
assert abs(Sample_Mean_MSE(-10,10)-3.33333333333333) < 0.05

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

## Problem 2

For this problem, you are asked to write a function **MinMax_Mean_MSE** that experimentally calculates the MSE of the min-max-mean estimator for samples drawn from uniform distribution given the bounds $a$ and $b$.

<font  style="color:blue"> **Code:**</font>
```python
MinMax_Mean_MSE(0,10)      
```


<font  style="color:magenta"> **Output**</font>
```
0.37865973633197908
```

```python
# modify this cell

def MinMax_Mean_MSE(a,b):
    # inputs: bounds for uniform distribution a and b
    # sample size is 10
    # number of experiments is 1,000,000
    # output: MSE for sample mean estimator with sample size 10
    
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: ex2, points: 5, grade: true, solution: false -->
```python
# Check Function
assert abs(MinMax_Mean_MSE(-5,5)-0.37882324043882964) < 0.05
assert abs(MinMax_Mean_MSE(-10,10)-1.5151124438250361) < 0.05

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```